---
title       : Analysing the data
description : 
---
## Find the total value of the dataset

```yaml
type: NormalExercise
key: 108d33a707
lang: python
xp: 100
skills: 2
```
Now that you are familiar with the data, you can finally start exploring it in depth. 

The first business question.



`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```
df['____'].____()

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---
## <<<New Exercise>>>

```yaml
type: NormalExercise
key: efa8b20bb4
lang: python
xp: 100
skills: 2
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
