---
title       : Exploring your data
description : Before any serious analysis can begin, it is important to understand 
---
## <<<New Exercise>>>

```yaml
type: NormalExercise
key: 5f798c6519
lang: python
xp: 100
skills: 2
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
