---
title       : Explore the data
description : Before any serious analysis can begin, it is important to understand your data. We do this by perfomring Exploratory Data Anlysis (EDA). A few of EDA methods will be covered in this chapter.
---
## Let's get some info! 

```yaml
type: NormalExercise
key: 5f798c6519
lang: python
xp: 100
skills: 2
```
We will begin by getting some very braod infromation about the Data Frame df. 

`@instructions`

Use the `info()` method on the df variable

`@pre_exercise_code`
```
import pandas as pd 

```

`@sample_code`
```
df.____()

```

`@solution`
```
df.info()
```

`@sct`
```{python}

```
