---
title: Getting started
description: In this chapter you will learn how 

---
## <<<New Exercise>>>

```yaml
type: MultipleChoiceExercise
key: d4540636bd
lang: python
xp: 50
skills: 2
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```
---
## Import the required libraries

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 2
key: 286ddf980c
```


`@instructions`

`@hint`

`@pre_exercise_code`
```
```

`@sample_code`
```
#Import numpy library

import numpy as np

#Import pandas library

import ____ as ____

```

`@solution`
```
#Import numpy library

import numpy as np

#Import pandas library

import pandas as pd

```

`@sct`
```{python}

```


---
## Import the data

```yaml
type: NormalExercise
key: 56c6ebec11
lang: python
xp: 100
skills: 2
```


`@instructions`

`@hint`

`@pre_exercise_code`
```
import pandas as pd

```

`@sample_code`
```
df 

```

`@solution`
```{python}

```

`@sct`
```{python}

```
